<head> 
<!-- Javascript -->
<script type="text/javascript" src="static/js/runestonebase.js"></script>

<script type="text/javascript" src="static/js/jquery.js"></script>
<script type="text/javascript" src="static/js/jquery-ui-1.10.3.custom.min.js"></script>
<script src="static/skulpt/dist/skulpt.min.js" type="text/javascript"></script> 
<script src="static/skulpt/dist/skulpt-stdlib.js" type="text/javascript"></script>
<script src="/static/js/javascript.js" type="text/javascript"></script>

<script src="static/codemirror/lib/codemirror.js" type="text/javascript"></script>
<script src="/static/js/cm-validator-remote.js" type="text/javascript"></script>

<script src="/static/codemirror/lib/lint.js""></script>
<script src="static/codemirror/lib/python.js" type="text/javascript"></script>

<script src="/static/codemirror/addon/dialog.js""></script>

<script src="static/js/activecode.js" type="text/javascript"></script>
<script src="static/js/timed_activecode.js" type="text/javascript"></script>

<script src="static/bootstrap-3.0.0/js/bootstrap.min.js" type="text/javascript"></script>

<!--Styling -->
<link href='https://fonts.googleapis.com/css?family=Raleway' rel='stylesheet' type='text/css'>
<link href='https://fonts.googleapis.com/css?family=Inconsolata' rel='stylesheet' type='text/css'>
<link rel="stylesheet" href="static/codemirror/lib/codemirror.css">
<link rel="stylesheet" href="static/css/activecode.css">

<link rel="stylesheet" href="static/bootstrap-3.0.0/css/bootstrap.min.css" type="text/css" />
<link rel="stylesheet" href="static/bootstrap-3.0.0/css/bootstrap-switch.min.css" type="text/css" />

<link rel="stylesheet" href="static/css/jquery-ui-1.10.3.custom.min.css" type="text/css" />
<link rel="stylesheet" href="static/css/custom_bootstrap.css" type="text/css" />

<link rel="stylesheet" href="/static/codemirror/lib/lint.css">
<link rel="stylesheet" href="/static/codemirror/addon/dialog.css">

<!-- OpenGraph Protocol -->

<meta property='og:image' content='static/image/PythonHelper.png'/>

<meta name=viewport content="width=device-width, initial-scale=1">
<!-- Config For Rune-->
<script type="text/javascript">
  eBookConfig = {};
  eBookConfig.useRunestoneServices = true;
</script>

<title>Python Helper</title>
</head> 
<body>
<a href="https://github.com/ethanchewy/PythonHelper"><img style="position: absolute; top: 0; right: 0; border: 0;" src="https://camo.githubusercontent.com/a6677b08c955af8400f44c6298f40e7d19cc5b2d/68747470733a2f2f73332e616d617a6f6e6177732e636f6d2f6769746875622f726962626f6e732f666f726b6d655f72696768745f677261795f3664366436642e706e67" alt="Fork me on GitHub" data-canonical-src="https://s3.amazonaws.com/github/ribbons/forkme_right_gray_6d6d6d.png"></a>

<div>
    <div class="text-center" id ="info">
    <h1>Python Linter Online</h1>
    <i>Live Syntax Checking Using Pylint</i>
    <p>
        Examples:
        <a href="#" id="codeexample1">1</a>
        <a href="#" id="codeexample2">2</a>
        <a href="#" id="codeexample3">3</a>
        <a href="#" id="codeexample5">5</a>
        <a href="#" id="codeexample6">6</a>
        <a href="#" id="codeexample7">7</a>
    </p>
</div>

    <!-- Text Editor -->
    <div class="row" id="core">
        <textarea id="txt" data-component="activecode" id=ac_example1 data-lang="python"    data-timelimit=25000  data-codelens="false" data-audio=''     data-caption='factorial'>
def factorial(n):
    return 1 if (n < 0) else n * factorial(n-1)sdf

====
from unittest.gui import TestCaseGui

class myTests(TestCaseGui):

    def testOne(self):
        self.assertEqual(factorial(3),6,"Try checking base case")
        self.assertEqual(factorial(0),1,"Try using recursion")

myTests().main()
        </textarea>
    </div>
    <div id="output">
        <pre></pre>
    </div>
    <!-- Grader -->
  <div id="grader-content">
    <div id="grader"></div>
  </div>
  <div id="append_text">
        <h2>Pylint Output Info</h2>
        <table id="errorslist" style="display: inline-block;">
            <tr>
                <th>Line</th>
                <th>Severity</th>
                <th>Error</th>
                <th>More Info</th>
            </tr>
        </table>
    </div>

  <footer>
    Created by Ethan Chiu. Source Code: <a href="https://github.com/ethanchewy/PythonBuddy">https://github.com/ethanchewy/PythonBuddy</a>
    </footer>
</body>

</html> 