<!DOCTYPE HTML>
<html>
<head>
    <meta http-equiv="Content-type" content="text/html; charset=utf-8"/>
    <meta http-equiv="Content-Language" content="ru-ru"/>
    <meta name="keywords" content="pyflakes, online checker, tools, python" />
    <meta name="description" content="Pyflakes Ajax Text Area">

    <!-- Scripts -->
    <script src="/static/js/jquery.js" type="text/javascript"></script>

    <script type="text/javascript" src="static/js/runestonebase.js"></script>
    <script src="static/js/activecode.js" type="text/javascript"></script>
    <script src="static/js/timed_activecode.js" type="text/javascript"></script>

    <script src="/static/codemirror/lib/codemirror.js"></script>
    <script src="/static/codemirror/lib/python.js"></script>
    <script src="/static/codemirror/lib/lint.js""></script>
    <script src="/static/js/cm-validator-remote.js" type="text/javascript"></script>

    
    <!-- Config For Rune-->
    <script type="text/javascript">
      eBookConfig = {};
      eBookConfig.useRunestoneServices = true;
    </script>

    <!--Codemirror Addons-->
    <script src="/static/codemirror/addon/search.js""></script>
    <script src="/static/codemirror/addon/searchcursor.js""></script>
    <script src="/static/codemirror/addon/dialog.js""></script>
    <!-- Skulpt -->
    <script src="/static/js/skulpt.min.js" type="text/javascript"></script> 
    <script src="/static/js/skulpt-stdlib.js" type="text/javascript"></script> 
    <!--Style -->
    
    <link rel="stylesheet" href="/static/css/bootstrap.min.css" type="text/css">

    <link rel="stylesheet" href="/static/css/style.css" type="text/css">

    <link rel="stylesheet" href="/static/codemirror/lib/codemirror.css">

    <link rel="stylesheet" href="/static/codemirror/lib/lint.css">

    <link rel="stylesheet" href="/static/codemirror/addon/dialog.css">


</head>
<body>
<!--<textarea data-lang="python" id="txt">print "hello world"</textarea>-->
<div class="text-center" id ="info">
    <h1>Python Linter Online</h1>
    <i>Live Syntax Checking Using Pylint</i>
    <p>
        Examples:
        <a href="#" id="codeexample1">1</a>
        <a href="#" id="codeexample2">2</a>
        <a href="#" id="codeexample3">3</a>
        <a href="#" id="codeexample5">5</a>
        <a href="#" id="codeexample6">6</a>
        <a href="#" id="codeexample7">7</a>
    </p>
</div>
<br>
<div class="container">
    <div class="row" id="core">
      <div class="col-md-8">
        <textarea id="txt" data-component="activecode" id=ac_example1 data-lang="python" data-timelimit=25000 data-codelens="false" data-audio='' data-caption='factorial'>
def factorial(n):
  return 1 if (n < 0) else n * factorial(n-1)

====
from unittest.gui import TestCaseGui

class myTests(TestCaseGui):

    def testOne(self):
        self.assertEqual(factorial(3),6,"Try checking base case")
        self.assertEqual(factorial(0),1,"Try using recursion")

myTests().main()
        </textarea>
      </div>
    </div>
    <div class="col-md-4" id="grader_container">
      <h3>Grader <i>(In Progress) </i><input id="grader_check" type="checkbox" name="my-checkbox" size="normal" data-on-color="success" data-on-color="info" data-on-text="Show"  data-off-text="Close" checked></h3>
      <div id="grader-content">
      </div>
    </div>
    <br>
    <button type="button" id="run">Run</button> 
    <i>Run function still in development. Note that all errors have to be resovlved for the program to run.</i>
    <pre id="output" ></pre> 
    <br>
    <div id="append_text">
        <h2>Pylint Output Info</h2>
        <table id="errorslist" style="display: inline-block;">
            <tr>
                <th>Line</th>
                <th>Severity</th>
                <th>Error</th>
                <th>More Info</th>
            </tr>
        </table>
    </div>
</div>

<br>
<footer>
Created by Ethan Chiu. Source Code: <a href="https://github.com/ethanchewy/PythonBuddy">https://github.com/ethanchewy/PythonBuddy</a>
</footer>
</body>
<script src="/static/js/javascript.js" type="text/javascript"></script>
</html>